<img src='https://hammondm.github.io/hltlogo1.png' style="float:right">

Linguistics 478/578<br>
Fall 2025<br>
Hammond

## Things to remember about any homework assignment:

1. For this assignment, you will edit this jupyter notebook and turn it in. Do not turn in pdf files or separate `.py` files.
1. Late work is not accepted.
1. Given the way I grade, you should try to answer *every* question, even if you don't like your answer or have to guess.
1. You may *not* use `python` modules that we have not already used in class.
1. You may certainly talk to your classmates about the assignment, but everybody must turn in *their own* work. It is not acceptable to turn in work that is essentially the same as the work of classmates.
1. All code must run. It doesn't have to be perfect, it may not do all that you want it to do, but it must run without error.
1. Code must run in reasonable time. Assume that if it takes more than *5 minutes* to run (on your machine), that's too long.
1. Make sure to remove or comment out the `raise NotImplementedError()` line in your answers.
1. Please do not add, remove, or copy autograded cells.
1. Do *not* try to do this assignment on Colab or with VSCode. Those will break the notebook and you'll get no points.
1. Make sure to select `restart, run all cells` from the `kernel` menu when you're done and before you turn this in!

## Homework #1

Here are the imports. Please do not import anything else.

```python
import numpy as np
from scipy.fftpack import rfft
```

0. Enter your name below.

<!-- grade_id: q0a, points: 0, grade: false, solution: true -->
```python
#definte myname
#myname = "?"
# YOUR CODE HERE, Define the function
myname = "Davo Acevedo-Cardona"
```

<!-- grade_id: q0t1, points: 1, grade: true, solution: false -->
```python
assert len(myname.split()) > 1
```

1. If we calculate sine from $0$ to $60$, how many cycles is that? (Be precise to two decimal points.)

<!-- grade_id: q1a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
#define x
x = 60/360 #(360 because is equal to 2pi and I'm assuming we are talking about degrees°)
```

<!-- grade_id: q1t1, points: 1, grade: true, solution: false -->
```python
#this test is hidden!
```

2. Write a function that will return a sine wave with `freq` cycles over `samples` samples.

<!-- grade_id: q2a, points: 0, grade: false, solution: true -->
```python
def makesine(freq,samples):
    '''make a sine wave
    args:
        freq: how many cycles
        samples: how many samples
    returns:
        a numpy array with the sine wave
    '''
    # YOUR CODE HERE
    #First and foremost I'm going to create the time vector (Y-AXIS):
    t = np.linspace(0, 1, samples)
    
    #Return the sine wave
    return np.sin(2 * np.pi * freq * t)
```

<!-- grade_id: q2t1, points: 1, grade: true, solution: false -->
```python
res = makesine(10,1000)
assert len(res) == 1000
```

<!-- grade_id: q2t2, points: 1, grade: true, solution: false -->
```python
assert np.isclose(max(res),1,atol=.1)
```

<!-- grade_id: q2t3, points: 1, grade: true, solution: false -->
```python
assert np.isclose(min(res),-1,atol=.1)
```

<!-- grade_id: q2t4, points: 1, grade: true, solution: false -->
```python
zeros = 0
for i in range(1,len(res)):
    if res[i]*res[i-1] <= 0: zeros += 1
zeros //= 2
assert zeros == 10
```

3. Now write a function that will return a sine wave with a specific frequency, sample rate, *and* duration (in seconds). (You may call your previous function or not, as you choose.)

<!-- grade_id: q3a, points: 0, grade: false, solution: true -->
```python
def makewave(freq,samples,duration):
    '''make a sine wave with a specific duration
    args:
        freq: frequency of the wave
        samples: how many samples per second
        duration: how long (in seconds)
    returns:
        the wave as a numpy array
    '''
    # YOUR CODE HERE
    # create the vector for samples:
    sam = int(samples * duration)
    
    # Now in this case the time vector will vary between 0 until the duration of the sample
    t = np.linspace(0, duration, sam)
    
    # retunr the sine wave
    return np.sin(2 * np.pi * freq * t)
```

<!-- grade_id: q3t1, points: 1, grade: true, solution: false -->
```python
res2 = makewave(100,1000,1.2)
assert len(res2) == 1200
```

<!-- grade_id: q3t2, points: 1, grade: true, solution: false -->
```python
assert np.isclose(max(res2),1,atol=.1)
```

<!-- grade_id: q3t3, points: 1, grade: true, solution: false -->
```python
assert np.isclose(min(res2),-1,atol=.1)
```

<!-- grade_id: q3t4, points: 1, grade: true, solution: false -->
```python
zeros = 0
for i in range(1,len(res2)):
    if res2[i]*res2[i-1] <= 0: zeros += 1
zeros //= 2
assert zeros == 120
```

4. Now write a function that returns whether some hypothetical harmonic is consistent with a specific fundamental frequency.

<!-- grade_id: q4a, points: 0, grade: false, solution: true -->
```python
def harmonic(f0,h):
    '''checks if a harmonic frequency is legit
    args:
        f0: fundamental frequency of the complex wave
        h: harmonic frequency
    returns:
        boolean
    '''
    # YOUR CODE HERE
    # This part calculate which is the harmonic of the signal
    n = h/f0
    
    # check if harmonic is a integer
    return np.isclose(n, round(n)) #Rounds to the closest int. (used for case 3)
```

<!-- grade_id: q4t1, points: 1, grade: true, solution: false -->
```python
assert harmonic(50,150)
```

<!-- grade_id: q4t2, points: 1, grade: true, solution: false -->
```python
assert harmonic(50,50)
```

<!-- grade_id: q4t3, points: 1, grade: true, solution: false -->
```python
assert harmonic(60,100) == False
```

5. (**578 only**) Write a function that will generate a pulse train by combining harmonics. Your function will take four arguments: the fundamental frequency, the maximum frequency for the harmonics, the duration (in seconds), and the sample rate (per second). All harmonics should have the same amplitude and the maximum amplitude of the complex wave should be $1$. (You may call your previous functions if you choose.)

<!-- grade_id: q5a, points: 0, grade: false, solution: true -->
```python
def pulsetrain(f0,freqmax,duration,samples):
    '''make a pulse train
    args:
        f0: fundamental frequency of the complex wave
        freqmax: maximum frequency for harmonics
        duration: duration in seconds
        samples: sample rate (per second)
    returns:
        wave as numpy array
    '''
    # YOUR CODE HERE
    # create the vector for samples and time (Same as Question 3):
    sam = int(samples * duration)
    t = np.linspace(0, duration, sam)
    
    # Create the sine waves:
    signal = np.zeros(sam) #Array of zeros
    
    # determine harmonics as an int
    h = int(freqmax/f0)
    
    #Fundamental and harmonics
    for n in range(1, h + 1): #"for" cycle to fill the array
        signal = signal + np.sin(2*np.pi*f0*t*n)
    
    # normalize amplitude to 1
    if np.max(np.abs(signal)) > 0:
        signal = signal / np.max(np.abs(signal))
    
    return signal
```

<!-- grade_id: q5t1, points: 1, grade: true, solution: false -->
```python
res3 = pulsetrain(100,2000,1,10000)
assert len(res3) == 10000
```

<!-- grade_id: q5t2, points: 1, grade: true, solution: false -->
```python
assert np.isclose(res3.max(),1.0,atol=.1)
```

<!-- grade_id: q5t3, points: 1, grade: true, solution: false -->
```python
yf = rfft(res3)
assert yf[800] < -250
```

<!-- grade_id: q5t4, points: 1, grade: true, solution: false -->
```python
assert yf[3800] < -250
```